[1434306966] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.9% (282332 kB) free!
[1434307565] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.7% (134616 kB) free!
[1434311164] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.1% (239896 kB) free!
[1434313564] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.9% (186188 kB) free!
[1434319565] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 15.9% (377300 kB) free!
[1434320164] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.3% (125452 kB) free!
[1434326165] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 16.7% (396324 kB) free!
[1434326764] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.5% (129312 kB) free!
[1434338163] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1434351748] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;1;connect to address 103.255.101.13 and port 25: Connection refused
[1434351808] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 25: Connection refused
[1434351868] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;3;connect to address 103.255.101.13 and port 25: Connection refused
[1434351928] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;HARD;4;connect to address 103.255.101.13 and port 25: Connection refused
[1434352229] SERVICE ALERT: ns3.tld6.net;SMTP;OK;HARD;4;SMTP OK - 0.666 sec. response time
[1434353163] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 29.5% (697544 kB) free.
[1434353762] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;1;CRITICAL - 6.5% (154412 kB) free!
[1434353882] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;2;CRITICAL - 6.8% (160112 kB) free!
[1434354003] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.4% (128812 kB) free!
[1434356361] SERVICE ALERT: ns3.tld6.net;SWAP;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1434356361] SERVICE ALERT: ns3.tld6.net;IMAP;CRITICAL;SOFT;1;connect to address 103.255.101.13 and port 143: Connection refused
[1434356438] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434356449] SERVICE ALERT: ns3.tld6.net;USER;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1434356460] HOST ALERT: ns3.tld6.net;DOWN;SOFT;1;CRITICAL - Host Unreachable (103.255.101.13)
[1434356481] SERVICE ALERT: ns3.tld6.net;SWAP;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1434356481] SERVICE ALERT: ns3.tld6.net;IMAP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 143: Connection refused
[1434356485] HOST ALERT: ns3.tld6.net;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1434356488] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 25: Connection refused
[1434356549] SERVICE ALERT: ns3.tld6.net;SMTP;OK;SOFT;3;SMTP OK - 0.943 sec. response time
[1434356559] SERVICE ALERT: ns3.tld6.net;USER;OK;SOFT;2;USERS OK - 1 users currently logged in
[1434356601] SERVICE ALERT: ns3.tld6.net;SWAP;OK;SOFT;3;SWAP OK - 100% free (15359 MB out of 15359 MB)
[1434356601] SERVICE ALERT: ns3.tld6.net;IMAP;OK;SOFT;3;IMAP OK - 0.001 second response time on 103.255.101.13 port 143 [* OK [CAPABILITY IMAP4rev1 LITERAL+ SASL-IR LOGIN-REFERRALS ID ENABLE IDLE NAMESPACE STARTTLS AUTH=PLAIN AUTH=LOGIN] Dovecot ready.]
[1434357602] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 13.2% (520996 kB) free!
[1434358202] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 0.7% (28408 kB) free!
[1434358202] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1434369602] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)