[1434669758] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 12.3% (486684 kB) free!
[1434672158] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.6% (301288 kB) free!
[1434696758] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 24.1% (953936 kB) free.
[1434697359] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;1;CRITICAL - 0.9% (35244 kB) free!
[1434697478] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;2;CRITICAL - 3.6% (142920 kB) free!
[1434697598] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 0.5% (18340 kB) free!
[1434712427] SERVICE ALERT: ns3.tld6.net;USER;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1434712547] SERVICE ALERT: ns3.tld6.net;USER;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1434712667] SERVICE ALERT: ns3.tld6.net;USER;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1434716867] SERVICE ALERT: ns3.tld6.net;USER;OK;HARD;3;USERS OK - 5 users currently logged in
[1434716986] SERVICE ALERT: wns1.tld6.net;Explorer;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434716989] HOST ALERT: wns1.tld6.net;DOWN;SOFT;1;CRITICAL - Host Unreachable (103.255.101.27)
[1434717070] HOST ALERT: wns1.tld6.net;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.18 ms
[1434717106] SERVICE ALERT: wns1.tld6.net;Explorer;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1434717216] SERVICE ALERT: wns1.tld6.net;Explorer;CRITICAL;HARD;3;Explorer.exe: not running
[1434717816] SERVICE ALERT: wns1.tld6.net;Explorer;OK;HARD;3;explorer.exe: Running
[1434718031] SERVICE ALERT: wns1.tld6.net;W3SVC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434718101] SERVICE ALERT: wns1.tld6.net;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434718141] SERVICE ALERT: wns1.tld6.net;W3SVC;OK;SOFT;2;W3SVC: Started
[1434718211] SERVICE ALERT: wns1.tld6.net;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1434720166] SERVICE ALERT: wns1.tld6.net;Memory Usage;CRITICAL;SOFT;1;connect to address 103.255.101.27 and port 12489: Connection refused
[1434720226] SERVICE ALERT: wns1.tld6.net;Explorer;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434720257] SERVICE ALERT: wns1.tld6.net;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434720288] SERVICE ALERT: wns1.tld6.net;Memory Usage;OK;SOFT;2;Memory usage: total:3583.64 MB - used: 1465.69 MB (41%) - free: 2117.96 MB (59%)
[1434720336] SERVICE ALERT: wns1.tld6.net;Explorer;OK;SOFT;2;explorer.exe: Running
[1434720367] SERVICE ALERT: wns1.tld6.net;NSClient++ Version;OK;SOFT;2;NSClient++ 0,4,1,73 2012-12-17