[1435015894] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 23.8% (944040 kB) free.
[1435017694] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;1;WARNING - 19.6% (775392 kB) free!
[1435017814] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;2;WARNING - 19.4% (768904 kB) free!
[1435017934] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 19.0% (750804 kB) free!
[1435021529] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.4% (293576 kB) free!
[1435041928] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.1% (398288 kB) free!
[1435042527] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.8% (267952 kB) free!
[1435062326] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 28.7% (1137864 kB) free.
[1435063427] SERVICE ALERT: wns1.tld6.net;Memory Usage;WARNING;SOFT;1;Memory usage: total:3583.64 MB - used: 2877.46 MB (80%) - free: 706.19 MB (20%)
[1435063493] SERVICE ALERT: wns1.tld6.net;Explorer;OK;HARD;3;explorer.exe: Running
[1435063547] SERVICE ALERT: wns1.tld6.net;Memory Usage;WARNING;SOFT;2;Memory usage: total:3583.64 MB - used: 2872.12 MB (80%) - free: 711.52 MB (20%)
[1435063667] SERVICE ALERT: wns1.tld6.net;Memory Usage;WARNING;HARD;3;Memory usage: total:3583.64 MB - used: 2870.97 MB (80%) - free: 712.68 MB (20%)
[1435064125] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;1;WARNING - 19.6% (777848 kB) free!
[1435064245] SERVICE ALERT: ns3.tld6.net;RAM;OK;SOFT;2;OK - 22.5% (890136 kB) free.
[1435064267] SERVICE ALERT: wns1.tld6.net;Memory Usage;OK;HARD;3;Memory usage: total:3583.64 MB - used: 2812.97 MB (78%) - free: 770.67 MB (22%)
[1435064845] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;1;WARNING - 17.6% (696596 kB) free!
[1435064965] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;2;WARNING - 17.0% (673164 kB) free!
[1435065085] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 15.7% (622504 kB) free!
[1435066936] SERVICE ALERT: ns3.tld6.net;SMTP;WARNING;SOFT;1;Invalid SMTP response received from host: 421 Unexpected log failure, please try later
[1435066996] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 25: Connection refused
[1435067056] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;3;connect to address 103.255.101.13 and port 25: Connection refused
[1435067116] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;HARD;4;connect to address 103.255.101.13 and port 25: Connection refused
[1435069013] SERVICE ALERT: ns3.tld6.net;POP3;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1435069133] SERVICE ALERT: ns3.tld6.net;POP3;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1435069228] SERVICE ALERT: ns3.tld6.net;IMAP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1435069253] SERVICE ALERT: ns3.tld6.net;POP3;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1435069348] SERVICE ALERT: ns3.tld6.net;IMAP;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1435069468] SERVICE ALERT: ns3.tld6.net;IMAP;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1435069885] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 9.6% (379884 kB) free!
[1435071625] HOST ALERT: ns3.tld6.net;DOWN;SOFT;1;CRITICAL - Host Unreachable (103.255.101.13)
[1435071637] SERVICE ALERT: ns3.tld6.net;SSH;CRITICAL;SOFT;1;connect to address 103.255.101.13 and port 2040: No route to host
[1435071639] HOST ALERT: ns3.tld6.net;DOWN;SOFT;2;CRITICAL - Host Unreachable (103.255.101.13)
[1435071691] HOST ALERT: ns3.tld6.net;DOWN;SOFT;3;CRITICAL - Host Unreachable (103.255.101.13)
[1435071754] SERVICE ALERT: ns3.tld6.net;SSH;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 2040: No route to host
[1435071754] HOST ALERT: ns3.tld6.net;DOWN;SOFT;4;CRITICAL - Host Unreachable (103.255.101.13)
[1435071780] SERVICE ALERT: ns3.tld6.net;FTP;CRITICAL;SOFT;1;connect to address 103.255.101.13 and port 21: No route to host
[1435071783] HOST ALERT: ns3.tld6.net;DOWN;SOFT;5;CRITICAL - Host Unreachable (103.255.101.13)
[1435071840] SERVICE ALERT: ns3.tld6.net;FTP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 21: No route to host
[1435071843] HOST ALERT: ns3.tld6.net;DOWN;SOFT;6;CRITICAL - Host Unreachable (103.255.101.13)
[1435071864] HOST ALERT: ns3.tld6.net;DOWN;SOFT;7;CRITICAL - Host Unreachable (103.255.101.13)
[1435071876] SERVICE ALERT: ns3.tld6.net;SSH;CRITICAL;HARD;3;connect to address 103.255.101.13 and port 2040: No route to host
[1435071879] HOST ALERT: ns3.tld6.net;DOWN;SOFT;8;CRITICAL - Host Unreachable (103.255.101.13)
[1435071884] SERVICE ALERT: ns3.tld6.net;HTTP;CRITICAL;SOFT;1;connect to address 103.255.101.13 and port 80: No route to host
[1435071900] SERVICE ALERT: ns3.tld6.net;FTP;CRITICAL;SOFT;3;connect to address 103.255.101.13 and port 21: No route to host
[1435071903] HOST ALERT: ns3.tld6.net;DOWN;SOFT;9;CRITICAL - Host Unreachable (103.255.101.13)
[1435071920] SERVICE ALERT: ns3.tld6.net;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (103.255.101.13)
[1435071920] HOST ALERT: ns3.tld6.net;DOWN;HARD;10;CRITICAL - Host Unreachable (103.255.101.13)
[1435071944] SERVICE ALERT: ns3.tld6.net;HTTP;CRITICAL;SOFT;2;connect to address 103.255.101.13 and port 80: No route to host
[1435071959] SERVICE ALERT: ns3.tld6.net;FTP;CRITICAL;HARD;4;connect to address 103.255.101.13 and port 21: No route to host
[1435071988] SERVICE ALERT: ns3.tld6.net;ZOMBIE;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435071996] SERVICE ALERT: ns3.tld6.net;CPU LOAD;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435072004] SERVICE ALERT: ns3.tld6.net;HTTP;CRITICAL;SOFT;3;connect to address 103.255.101.13 and port 80: No route to host
[1435072024] SERVICE ALERT: ns3.tld6.net;USER;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435072041] SERVICE ALERT: ns3.tld6.net;PING;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (103.255.101.13)
[1435072053] SERVICE ALERT: ns3.tld6.net;SWAP;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435072064] SERVICE ALERT: ns3.tld6.net;HTTP;CRITICAL;HARD;4;connect to address 103.255.101.13 and port 80: No route to host
[1435072081] SERVICE ALERT: ns3.tld6.net;DISK;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435072095] SERVICE ALERT: ns3.tld6.net;PROCESS;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1435072108] SERVICE ALERT: ns3.tld6.net;ZOMBIE;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072116] SERVICE ALERT: ns3.tld6.net;CPU LOAD;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072144] SERVICE ALERT: ns3.tld6.net;USER;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072161] SERVICE ALERT: ns3.tld6.net;PING;CRITICAL;HARD;3;CRITICAL - Host Unreachable (103.255.101.13)
[1435072173] SERVICE ALERT: ns3.tld6.net;SWAP;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072201] SERVICE ALERT: ns3.tld6.net;DISK;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072215] SERVICE ALERT: ns3.tld6.net;PROCESS;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1435072228] SERVICE ALERT: ns3.tld6.net;ZOMBIE;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435072236] SERVICE ALERT: ns3.tld6.net;CPU LOAD;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435072262] SERVICE ALERT: ns3.tld6.net;USER;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435072291] SERVICE ALERT: ns3.tld6.net;SWAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435072321] SERVICE ALERT: ns3.tld6.net;DISK;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435072335] SERVICE ALERT: ns3.tld6.net;PROCESS;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1435075062] SERVICE ALERT: ns3.tld6.net;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 316 bytes in 0.332 second response time
[1435075066] HOST ALERT: ns3.tld6.net;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1435075162] SERVICE ALERT: ns3.tld6.net;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1435075216] SERVICE ALERT: ns3.tld6.net;SMTP;OK;HARD;4;SMTP OK - 0.344 sec. response time
[1435075225] SERVICE ALERT: ns3.tld6.net;ZOMBIE;OK;HARD;3;PROCS OK: 0 processes with STATE = Z
[1435075233] SERVICE ALERT: ns3.tld6.net;CPU LOAD;OK;HARD;3;OK - load average: 1.25, 0.62, 0.23
[1435075243] SERVICE ALERT: ns3.tld6.net;POP3;OK;HARD;3;POP OK - 0.001 second response time on 103.255.101.13 port 110 [+OK Dovecot ready.]
[1435075256] SERVICE ALERT: ns3.tld6.net;FTP;OK;HARD;4;FTP OK - 0.002 second response time on 103.255.101.13 port 21 [220---------- Welcome to Pure-FTPd [privsep] [TLS] ----------
[1435075261] SERVICE ALERT: ns3.tld6.net;USER;OK;HARD;3;USERS OK - 0 users currently logged in
[1435075286] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 45.8% (1815844 kB) free.
[1435075290] SERVICE ALERT: ns3.tld6.net;SWAP;OK;HARD;3;SWAP OK - 100% free (15359 MB out of 15359 MB)
[1435075318] SERVICE ALERT: ns3.tld6.net;DISK;OK;HARD;3;DISK OK - free space: / 344054 MB (71% inode=79%):
[1435075332] SERVICE ALERT: ns3.tld6.net;PROCESS;OK;HARD;3;PROCS OK: 4 processes
[1435075456] SERVICE ALERT: ns3.tld6.net;IMAP;OK;HARD;3;IMAP OK - 0.001 second response time on 103.255.101.13 port 143 [* OK [CAPABILITY IMAP4rev1 LITERAL+ SASL-IR LOGIN-REFERRALS ID ENABLE IDLE NAMESPACE STARTTLS AUTH=PLAIN AUTH=LOGIN] Dovecot ready.]
[1435075473] SERVICE ALERT: ns3.tld6.net;SSH;OK;HARD;3;SSH OK - OpenSSH_5.3 (protocol 2.0)
[1435078885] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;1;WARNING - 14.8% (588104 kB) free!
[1435079005] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;SOFT;2;WARNING - 15.3% (606856 kB) free!
[1435079125] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 18.1% (718736 kB) free!
[1435080325] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 9.6% (379416 kB) free!
[1435080325] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1435080925] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 12.6% (497464 kB) free!
[1435081525] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 8.7% (346520 kB) free!
[1435082725] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.5% (414536 kB) free!
[1435083925] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.6% (299788 kB) free!