[1434140168] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1434143168] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.2% (266188 kB) free!
[1434144368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.3% (148180 kB) free!
[1434150968] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.4% (245780 kB) free!
[1434151568] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.0% (118336 kB) free!
[1434151568] SERVICE FLAPPING ALERT: ns3.tld6.net;RAM;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1434152768] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.6% (274736 kB) free!
[1434153368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.8% (159848 kB) free!
[1434153968] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.4% (246708 kB) free!
[1434154568] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.1% (168264 kB) free!
[1434155768] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.1% (263824 kB) free!
[1434156368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 9.1% (215164 kB) free!
[1434159368] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.2% (241996 kB) free!
[1434159968] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.4% (176276 kB) free!
[1434161168] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.4% (245160 kB) free!
[1434162368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.4% (174528 kB) free!
[1434165968] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.3% (268316 kB) free!
[1434166568] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.6% (180460 kB) free!
[1434169568] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 15.1% (356976 kB) free!
[1434170168] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 8.7% (206416 kB) free!
[1434170768] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.0% (259640 kB) free!
[1434171368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 8.9% (211828 kB) free!
[1434173768] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.8% (279312 kB) free!
[1434174368] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.3% (124736 kB) free!
[1434176168] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.5% (249696 kB) free!
[1434176768] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.1% (121340 kB) free!
[1434177776] SERVICE ALERT: ns3.tld6.net;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434177826] SERVICE ALERT: ns3.tld6.net;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 316 bytes in 0.001 second response time
[1434178569] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.0% (261200 kB) free!
[1434178794] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434178847] SERVICE ALERT: ns3.tld6.net;SMTP;OK;SOFT;2;SMTP OK - 3.216 sec. response time
[1434179169] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.7% (135972 kB) free!
[1434180369] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 14.0% (330768 kB) free!
[1434180969] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.1% (121556 kB) free!
[1434181570] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 44.2% (1047056 kB) free.
[1434182170] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;1;CRITICAL - 7.7% (182948 kB) free!
[1434182288] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;2;CRITICAL - 5.1% (121464 kB) free!
[1434182408] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.8% (137848 kB) free!
[1434189008] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.8% (255564 kB) free!
[1434189609] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 7.5% (176764 kB) free!
[1434193253] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434193304] SERVICE ALERT: ns3.tld6.net;SMTP;OK;SOFT;2;SMTP OK - 1.076 sec. response time
[1434195006] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 12.3% (290444 kB) free!
[1434195413] SERVICE ALERT: ns3.tld6.net;SMTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1434195469] SERVICE ALERT: ns3.tld6.net;SMTP;OK;SOFT;2;SMTP OK - 5.989 sec. response time
[1434195606] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.2% (146888 kB) free!
[1434197406] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.5% (271848 kB) free!
[1434198006] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 8.1% (192280 kB) free!
[1434201006] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 11.1% (262768 kB) free!
[1434201606] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.7% (158112 kB) free!
[1434202806] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 10.3% (244656 kB) free!
[1434203406] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 5.9% (139904 kB) free!
[1434206406] SERVICE ALERT: ns3.tld6.net;RAM;WARNING;HARD;3;WARNING - 12.5% (295012 kB) free!
[1434207006] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 6.9% (163264 kB) free!
[1434209406] SERVICE ALERT: ns3.tld6.net;RAM;OK;HARD;3;OK - 24.8% (587876 kB) free.
[1434210006] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;1;CRITICAL - 5.1% (120032 kB) free!
[1434210126] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;SOFT;2;CRITICAL - 4.9% (116488 kB) free!
[1434210246] SERVICE ALERT: ns3.tld6.net;RAM;CRITICAL;HARD;3;CRITICAL - 8.5% (200064 kB) free!